Visual experience shapes functional connectivity between occipital and non-visual networks

Author response:

Reviewer #1 (Public Review):

Summary:

The present study evaluates the role of visual experience in shaping functional correlations between extrastriate visual cortex and frontal regions. The authors used fMRI to assess "resting-state" temporal correlations in three groups: sighted adults, congenitally blind adults, and neonates. Previous research has already demonstrated differences in functional correlations between visual and frontal regions in sighted compared to early blind individuals. The novel contribution of the current study lies in the inclusion of an infant dataset, which allows for an assessment of the developmental origins of these differences.

The main results of the study reveal that correlations between prefrontal and visual regions are more prominent in the blind and infant groups, with the blind group exhibiting greater lateralization. Conversely, correlations between visual and somato-motor cortices are more prominent in sighted adults. Based on these data, the authors conclude that visual experience plays an instructive role in shaping these cortical networks. This study provides valuable insights into the impact of visual experience on the development of functional connectivity in the brain.

Strengths:

The dissociations in functional correlations observed among the sighted adult, congenitally blind, and neonate groups provide strong support for the study's main conclusion regarding experience-driven changes in functional connectivity profiles between visual and frontal regions.

In general, the findings in sighted adult and congenitally blind groups replicate previous studies and enhance the confidence in the reliability and robustness of the current results.

Split-half analysis provides a good measure of robustness in the infant data.

Weaknesses:

There is some ambiguity in determining which aspects of these networks are shaped by experience.

This uncertainty is compounded by notable differences in data acquisition and preprocessing methods, which could result in varying signal quality across groups. Variations in signal quality may, in turn, have an impact on the observed correlation patterns.

The study's findings could benefit from being situated within a broader debate surrounding the instructive versus permissive roles of experience in the development of visual circuits.

Reviewer #2 (Public Review):

Summary:

Tian et al. explore the developmental organs of cortical reorganization in blindness. Previous work has found that a set of regions in the occipital cortex show different functional responses and patterns of functional correlations in blind vs. sighted adults. In this paper, Tian et al. ask: how does this organization arise over development? Is the "starting state" more like the blind pattern, or more like the adult pattern? Their analyses reveal that the answer depends on the particular networks investigated; some functional connections in infants look more like blind than sighted adults; other functional connections look more like sighted than blind adults; and others fall somewhere in the middle, or show an altogether different pattern in infants compared with both sighted and blind adults.

Strengths:

The question raised in this paper is extremely important: what is the starting state in development for visual cortical regions, and how is this organization shaped by experience? This paper is among the first to examine this question, particularly by comparing infants not only with sighted adults but also blind adults, which sheds new light on the role of visual (and cross-modal) experience. Another clear strength lies in the unequivocal nature of many results. Many results have very large effect sizes, critical interactions between regions and groups are tested and found, and infant analyses are replicated in split halves of the data.

Weaknesses:

A central claim is that "infant secondary visual cortices functionally resemble those of blind more than sighted adults" (abstract, last paragraph of intro). I see two potential issues with this claim. First, a minor change: given the approaches used here, no claims should be made about the "function" of these regions, but rather their "functional correlations". Second (and more importantly), the claim that the secondary visual cortex in general resembles blind more than sighted adults is still not fully supported by the data. In fact, this claim is only true for one aspect of secondary visual area functional correlations (i.e., their connectivity to A1/M1/S1 vs. PFC). In other analyses, the infant secondary visual cortex looks more like sighted adults than blind adults (i.e., in within vs. across hemisphere correlations), or shows a different pattern from both sighted and blind adults (i.e., in occipito-frontal subregion functional connectivity). It is not clear from the manuscript why the comparison to PFC vs. non-visual sensory cortex is more theoretically important than hemispheric changes or within-PFC correlations (in fact, if anything, the within-PFC correlations strike me as the most important for understanding the development and reorganization of these secondary visual regions). It seems then that a more accurate conclusion is that the secondary visual cortex shows a mix of instructive effects of vision and reorganizing effects of blindness, albeit to a different extent than the primary visual cortex.

Relatedly, group differences in overall secondary visual cortex connectivity are particularly striking as visualized in the connectivity matrices shown in Figure S1. In the results (lines 105-112), it is noted that while the infant FC matrix is strongly correlated with both adult groups, the infant group is nonetheless more strongly correlated with the blind than sighted adults. I am concerned that these results might be at least partially explained by distance (i.e., local spread of the bold signal), since a huge portion of the variance in these FC matrices is driven by stronger correlations between regions within the same system (e.g., secondary-secondary visual cortex, frontal-frontal cortex), which are inherently closer together, relative to those between different systems (e.g., visual to frontal cortex). How do results change if only comparisons between secondary visual regions and non-visual regions are included (i.e., just the pairs of regions within the bold black rectangle on the figure), which limits the analysis to long-rang connections only? Indeed, looking at the off-diagonal comparisons, it seems that in fact there are three altogether different patterns here in the three groups. Even if the correlation between the infant pattern and blind adult pattern survives, it might be more accurate to claim that infants are different from both adult groups, suggesting both instructive effects of vision and reorganizing effects of blindness. It might help to show the correlation between each group and itself (across independent sets of subjects) to better contextualize the relative strength of correlations between the groups.

It is not clear that differences between groups should be attributed to visual experience only. For example, despite the title of the paper, the authors note elsewhere that cross-modal experience might also drive changes between groups. Another factor, which I do not see discussed, is possible ongoing experience-independent maturation. The infants scanned are extremely young, only 2 weeks old. Although no effects of age are detected, it is possible that cortex is still undergoing experience-independent maturation at this very early stage of development. For example, consider Figure 2; perhaps V1 connectivity is not established at 2 weeks, but eventually achieves the adult pattern later in infancy or childhood. Further, consider the possibility that this same developmental progression would be found in infants and children born blind. In that case, the blind adult pattern may depend on blindness-related experience only (which may or may not reflect "visual" experience per se). To deal with these issues, the authors should add a discussion of the role of maturation vs. experience and temper claims about the role of visual experience specifically (particularly in the title).

The authors measure functional correlations in three very different groups of participants and find three different patterns of functional correlations. Although these three groups differ in critical, theoretically interesting ways (i.e., in age and visual/cross-modal experience), they also differ in many uninteresting ways, including at least the following: sampling rate (TR), scan duration, multi-band acceleration, denoising procedures (CompCor vs. ICA), head motion, ROI registration accuracy, and wakefulness (I assume the infants are asleep).

Addressing all of these issues is beyond the scope of this paper, but I do feel the authors should acknowledge these confounds and discuss the extent to which they are likely (or not) to explain their results. The authors would strengthen their conclusions with analyses directly comparing data quality between groups (e.g., measures of head motion and split-half reliability would be particularly effective).

Response #1: We appreciate the reviewer’s comments. In response, we have revised the paper to provide a more balanced summary of the data and clarified in the introduction which signatures the paper focuses on and why. Additionally, we have included several control analyses to account for other plausible explanations for the observed group differences. Specifically, we randomly split the infant dataset into two halves and performed split-half cross-validation. Across all comparisons, the results from the two halves were highly similar, suggesting that the effects are robust (see Supplementary Figures S3 and S4).

Furthermore, we compared the split-half noise ceiling across the groups (infants, sighted adults, and blind adults) and found no significant differences between them (details in response #6). Finally, we repeated our analysis after excluding infants with a radiology score of 4 or 5, and the results remained consistent, indicating that our findings are not confounded by potential brain anomalies (details in response #2).

We hope these control analyses help strengthen our conclusions.

Reviewer #3 (Public Review):

Summary:

This study aimed to investigate whether the differences observed in the organization of visual brain networks between blind and sighted adults result from a reorganization of an early functional architecture due to blindness, or whether the early architecture is immature at birth and requires visual experience to develop functional connections. This question was investigated through the comparison of 3 groups of subjects with resting-state functional MRI (rs-fMRI). Based on convincing analyses, the study suggests that: 1) secondary visual cortices showed higher connectivity to prefrontal cortical regions (PFC) than to non-visual sensory areas (S1/M1 and A1) in sighted infants like in blind adults, in contrast to sighted adults; 2) the V1 connectivity pattern of sighted infants lies between that of sighted adults (stronger functional connectivity with non-visual sensory areas than with PFC) and that of blind adults (stronger functional connectivity with PFC than with non-visual sensory areas); 3) the laterality of the connectivity patterns of sighted infants resembled those of sighted adults more than those of blind adults, but sighted infants showed a less differentiated fronto-occipital connectivity pattern than adults.

Strengths:

The question investigated in this article is important for understanding the mechanisms of plasticity during typical and impaired development, and the approach considered, which compares different groups of subjects including, neonates/infants and blind adults, is highly original.

-Overall, the analyses considered are solid and well-detailed. The results are quite convincing, even if the interpretation might need to be revised downwards, as factors other than visual experience may play a role in the development of functional connections with the visual system.

Weaknesses:

While it is informative to compare the "initial" state (close to birth) and the "final" states in blind and sighted adults to study the impact of post-natal and visual experience, this study does not analyze the chronology of this development and when the specialization of functional connections is completed. This would require investigating when experience-dependent mechanisms are important for the setting- establishment of multiple functional connections within the visual system. This could be achieved by analyzing different developmental periods in the same way, using open databases such as the Baby Connectome Project. Given the early, "condensed" maturation of the visual system after birth, we might expect sighted infants to show connectivity patterns similar to those of adults a few months after birth.

The rationale for mixing full-term neonates and preterm infants (scanned at term-equivalent age) from the dHCP 3rd release is not understandable since preterms might have a very different development related to prematurity and to post-natal (including visual) experience. Although the authors show that the difference between the connectivity of visual and other sensory regions, and the one of visual and PFC regions, do not depend on age at birth, they do not show that each connectivity pattern is not influenced by prematurity. Simply not considering the preterm infants would have made the analysis much more robust, and the full-term group in itself is already quite large compared with the two adult groups. The current study setting and the analyses performed do not seem to be an adequate and sufficient model to ascertain that "a few weeks of vision after birth is ... insufficient to influence connectivity".

In a similar way, excluding the few infants with detected brain anomalies (radiological scores higher or equal to 4) would strengthen the group homogeneity by focusing on infants supposed to have a rather typical neurodevelopment. The authors quote all infants as "sighted" but this is not guaranteed as no follow-up is provided.

Response #2: We appreciate the reviewer’s suggestion. We re-analyzed the infant cohort after excluding all cases with radiological scores ≥4 (n =39 infants excluded). The revised analysis confirmed that the connectivity patterns reported in the main text remain statistically unchanged (see Supplementary Fig. S11). This demonstrates the robustness of our findings to potential confounding effects from potential brain anomalies. We have explicitly clarified this in the revised Methods section (page 14, line 391in the manuscript).

In our dataset, newborns (average age at scan = 2.79 weeks) have very limited and immature vision. We agree with the reviewer that long-term visual outcomes cannot be guaranteed without follow-up data. The term "sighted infants" was used operationally to distinguish this cohort from congenitally blind populations.

The post-menstrual age (PMA) at scan of the infants is also not described. The methods indicate that all were scanned at "term-equivalent age" but does this mean that there is some PMA variability between 37 and 41 weeks? Connectivity measures might be influenced by such inter-individual variability in PMA, and this could be evaluated.

The rationale for presenting results on the connectivity of secondary visual cortices before one of the primary cortices (V1) was not clear to understand. Also, it might be relevant to better justify why only the connectivity of visual regions to non-visual sensory regions (S1-M1, A1) and prefrontal cortex (PFC) was considered in the analyses, and not the ones to other brain regions.

In relation to the question explored, it might be informative to reposition the study in relation to what others have shown about the developmental chronology of structural and functional long-distance and short-distance connections during pregnancy and the first postnatal months.

The authors acknowledge the methodological difficulties in defining regions of interest (ROIs) in infants in a similar way as adults. The reliability and the comparability of the ROIs positioning in infants is definitely an issue. Given that brain development is not homogeneous and synchronous across brain regions (in particular with the frontal and parietal lobes showing delayed growth), the newborn brain is not homothetic to the adult brain, which poses major problems for registration. The functional specialization of cortical regions is incomplete at birth. This raises the question of whether the findings of this study would be stable/robust if slightly larger or displaced regions had been considered, to cover with greater certainty the same areas as those considered in adults. And have other cortical parcellation approaches been considered to assess the ROIs robustness (e.g. MCRIB-S for full-terms)?

Recommendations for the Authors:

Reviewer #1(Recommendations for the authors):

Further consideration should be given to the underlying changes in network architecture that may account for differences in functional correlations across groups. An increase (or decrease) in correlation between two regions could signify an increase (decrease) in connection or communication between those regions. Alternatively, it might reflect an increase in communication or connection with a third region, while the physical connections/interactions between the two original regions remain unchanged. These possibilities lead to distinct mechanistic interpretations. For example, there are substantial changes in connectivity during early visual (e.g. Burkhalter A. 1993, Cerebral Cortex) and visuo-motor development (e.g., Csibra et al. 2000 Neuroreport). It's not clear whether increases in communication within the visual network and improvements in visuo-motor behavior (e.g., Yizhar et al. 2023 Frontiers in Neuroscience) wouldn't produce a qualitatively similar pattern of results.

Relatedly, the within-network correlation patterns between visual ROIs and frontal ROIs appear markedly different between sighted adults and infants (Supplementary Figure S1). To what extent do the differences in long-range correlations between visual and frontal regions reflect these within-network differences in functional organization?

Response #3: The reviewer is raising some interesting questions about possible mechanisms and network changes. Resting state studies are indeed always subject to possibility that some effects are mediated by a third, unobserved region. Prior whole-cortex connectivity analyses have observed primarily changes in occipito-frontal connectivity in blindness, so there is not a clear cortical ‘third region’ candidate (Deen et al., 2015). However, some thalamic affects have also been observed and could contribute to the phenomenon (Bedny et al., 2011). Resting state changes in correlation between two areas do not imply changes in strength of long-range anatomical connectivity. Indeed, in the current case they may well reflect differential functional coupling, rather than strengthening or weakening of anatomical connections. We now discuss this in the Discussion section on page 12, line 301 as follows:

“Despite these insights, many questions remain regarding the neurobiological mechanisms underlying experience-based functional connectivity changes and their relationship to anatomical development. Long-range anatomical connections between brain regions are already present in infants—even prenatally—though they remain immature (Huang et al., 2009; Kostović et al., 2019, 2021; Takahashi et al., 2012; Vasung, 2017). Functional connectivity changes may stem from local synaptic modifications within these stable structural pathways, consistent with findings that functional connectivity can vary independently of structural connection strength (Fotiadis et al., 2024). Moreover, functional connectivity has been shown to outperform structural connectivity in predicting individual behavioral differences, suggesting that experience-based functional changes may reflect finer-scale synaptic or network-level modulations not captured by macrostructural measures (Ooi et al., 2022). Prior studies also suggest that, even in adults, coordinated sensory-motor experience can lead to enhancement of functional connectivity across sensory-motor systems, indicating that large-scale changes in functional connectivity do not necessarily require corresponding changes in anatomical connectivity (Guerra-Carrillo et al., 2014; Li et al., 2018).”

It is not clear how changes in correlation patterns among visual areas would produce the connectivity between visual areas and prefrontal areas reported in the current study. Activity in visual areas drives correlations both among visual areas and between visual and prefrontal areas and the same is true of prefrontal corticies.

The findings from this study should be more closely linked to the extensive literature surrounding the debate on whether experience plays an instructive or permissive role in visual development (e.g., Crair 1999 Current Opin Neurobiol; Sur et al. 1999 J Neurobiol; Kiorpes 2016 J Neurosci; Stellwagen & Shatz 2002 Neuron; Roy et al. 2020 Nature Communications).

Response #4: The instructive role suggests that specific experiences or patterns of neural activity directly shape and organize neural circuitry, while the permissive role indicates that such experiences or activity merely enable other factors, such as molecular signals, to influence neural circuit formation(Crair, 1999; Sur et al., 1999). To distinguish whether experience plays an instructive or permissive role, it is essential to manipulate the pattern or information content of neural activity while maintaining a constant overall activity level (Crair, 1999; Roy et al., 2020; Stellwagen & Shatz, 2002). However, both the sighted and blind adult groups have had extensive experience and neural activity in the visual cortices. For the sighted group, activity in the visual cortex is partly driven by bottom-up input from the external environment, through the retina, LGN, and ultimately to the cortex. In contrast, the blind group’s visual cortex activity is partially driven by top-down input from non-visual networks. The precise role of this activity in shaping the observed connectivity patterns remains unclear. Although our study cannot speak to this issue directly, we now link to the relevant literature on page 12,line 320 of the manuscript in the Discussion section as follows:

“The current findings reveal both effects of vision and effects of blindness on the functional connectivity patterns of the visual cortex. A further open question is whether visual experience plays an instructive or permissive role in shaping neural connectivity patterns. An instructive role suggests that specific sensory experiences or patterns of neural activity directly shape and organize neural circuitry. In contrast, a permissive role implies that sensory experience or neural activity merely facilitates the influence of other factors—such as molecular signals—on the formation and organization of neural circuits (Crair, 1999; Sur et al., 1999). Studies with animals that manipulate the pattern or informational content of neural activity while keeping overall activity levels constant could distinguish between these hypotheses (Crair, 1999; Roy et al., 2020; Stellwagen & Shatz, 2002).”

The assertion that a few weeks of vision after birth is insufficient to influence connectivity is provocative. Though supported by the study's results, it would benefit from integration with research in animal models showing considerable malleability of networks from early experience (e.g., Akerman et al. 2002 Neuron; Li et al. 2006 Nature Neuroscience; Stacy et al. 2023 J Neuroscience).

Response #5: We thank the reviewer for their suggestion. The present study found that several weeks of postnatal visual experience is insufficient to significantly alter the long-term connectivity patterns of the visual cortices. While animal studies have shown that acute visual experience, or even exposure to visual stimuli through unopened eyelids, can robustly influence visual system development(Akerman et al., 2002; Li et al., 2008; Van Hooser et al., 2012). We think this discrepancy may be attributed to the substantial differences in developmental timelines between species. The human lifespan is much longer, and so is the human critical period, making it unclear how to map duration from one species to another. We briefly touched upon the time course issue in page 11 line 289 in the Discussion section as follows:

“The present results reveal the effects of experience on development of functional connectivity between infancy and adulthood, but do not speak to the precise time course of these effects. Infants in the current sample had between 0 and 20 weeks of visual experience. Comparisons across these infants suggests that several weeks of postnatal visual experience is insufficient to produce a sighted-adult connectivity profile. The time course of development could be anywhere between a few months and years and could be tested by examining data from children of different ages.”

Substantial differences between the groups are evident in several key aspects of the study, including the number of subjects, brain sizes, imaging parameters, and data preprocessing, all of which are likely to have an impact on the overall signal quality. To clarify how these differences might have impacted correlation differences between groups, it would be essential to include information on the noise ceilings for each correlation analysis within each group.

Response #6: We thank the reviewer for their suggestion. We now report the split-half noise ceiling for adult and infant groups. For each participant, we first split the rs-fMRI time series into two halves, then calculated the ROI-wise rsFC pattern from the two splits. The split-half noise ceiling was estimated according to Lage-Castellanos et al (2019). The noise ceilings of the three groups (infants: 0.90 ± 0.056,blind adults: 0.88 ± 0.041, sighted adults: 0.90 ± 0.055) showed no significant difference (One-way ANOVA, F(2,552) = 2.348, p = 0.097). Therefore, we believe that overall signal quality is unlikely to impact our results. We also add the relevant context in the Method section in page 16 Line 447 as follows:

“Substantial differences between the groups exist in this study, including the number of subjects, brain sizes, imaging parameters, and data preprocessing, all of which are likely to have an impact on the overall signal quality. To address this concern, we compared the split-half noise ceiling across the groups (infants, sighted adults, and blind adults). For each participant, we first split the rs-fMRI time series into two halves, then calculated the ROI-wise rsFC pattern from the two splits. The split-half noise ceiling was estimated according to Lage-Castellanos et al (Lage-Castellanos et al., 2019). The noise ceilings of the three groups (infants: 0.90 ± 0.056, blind adults: 0.88 ± 0.041, sighted adults: 0.90 ± 0.055) showed no significant difference (One-way ANOVA, F (2,552) = 2.348, p = 0.097). Therefore, overall signal quality is unlikely to impact our results.”

In general, it appears that the infant correlations are stronger compared to the other groups. While this could reflect increased coherence or lack of differentiation, it is also possible that it is simply due to the presence of a non-neuronal global signal. Such a signal has the potential to substantially limit the effective range of functional correlations and comparisons with adults. To address this, it is advisable to conduct control analyses aimed at assessing and potentially removing global signals.

Response #7: We agree with the reviewer that global signal regression (GSR) may help reduce non-neuronal artifacts, such as motion, cardiac, and respiratory signals, which are known to correlate with the global signal. However, the global signal also contains neural signals from gray matter, and removing it can introduce unwanted artifacts, especially for the current study. First, GSR can reduce the physiological accuracy of functional connectivity (FC); second, GSR may have differential effects across groups, potentially introducing additional artifacts in between-group comparisons, as noted by Murphy et al (Murphy & Fox, 2017). The CompCor method (Behzadi et al., 2007; Whitfield-Gabrieli & Nieto-Castanon, 2012) is capble to estimate the global non-neuronal artifacts like the GSR method. Meanwhile as it estimate global non-neuronal artifacts from signals within the white matter (WM) and cerebrospinal fluid (CSF) masks, but not the gray matter (GM), CompCor could introduce minimal unwanted bias to the GM signal.

Was there a difference in correlations for preterm vs term neonates? Recent research has suggested that preterm births can have an impact on functional networks, particularly in frontal cortices. e.g., Tokariev et al. 2019, Li et al. 2021 elife; Zhang et al. 2022 Fronteirs in Neuroscience.

Response #8: We have compared preterm and term neonates for all the main results, including the connectivity from the secondary visual cortex/V1 to non-visual sensory cortices versus prefrontal cortices, the laterality of occipito-frontal connectivity, and the specialization across different fronto-occipital networks. This information is reported in Page 6 line 169 and Supplementary Figure S7. The connectivities of full-term infants are generally higher than those of preterm infants. However, the connectivity patterns of term and preterm infants are very similar.

The consistency between the current results and prior work (e.g., Burton et al. 2014) is notable, particularly in the observed greater correlations in prefrontal regions and weaker correlations in somato-motor regions for early blind individuals compared to sighted. However, almost all visual-frontal correlations in both groups were negative in that prior study. Some discussion on why positive correlations were found in the current study could help to clarify.

Response #9: Many other papers have reported positive correlations similar to those found in our study (e.g., Deen et al., 2015; Kanjlia et al., 2021). In contrast, Burton's study identified predominantly negative visual-frontal correlations, we think this is likely because the global signal was regressed out during preprocessing. This methodological choice can lead to an increase in negative connections (Murphy & Fox, 2017).

The term "secondary visual areas" used throughout the paper lacks specificity, and its usage in terms of underlying anatomical and functional areas has been inconsistent in the literature. It would be advisable to adopt a more precise characterization based on functional and/or anatomical criteria.

Response #10: We specified in the article that Tthe occipital ROIs were defined in the current study are functional areas in people born blind identified in prior studies as regions that respond to three non-visual tasks such as language, math, or executive function, and show functional connectivity changes in blind adults in previous studies (Kanjlia et al., 2016, 2021; Lane et al., 2015). These regions respond to language, math and executivie function in the congenitally blind population (see Figure 1.) The are refered collectively as ‘secondary visual areas’ to destinguish them from V1. Anatomically, these three regions cover the majority of the lateral occipital cortex and part of the ventral occipital cortex, providing a good sample of the connectivity profile of higher-order visual areas. Thus, we are using the term "secondary visual areas" to refer to these regions. In blind individuals, although these regions respond to non-visual tasks, their exact functions are unknown.

The inclusion of the ventral temporal cortex in the visual ROIs is currently only depicted in Supplementary Figure S7. To enhance the clarity of the areas of interest analyzed, it would be advisable to illustrate the ventral temporal areas in the main text. Were there notable differences in the frontal correlations between the lateral occipital visual areas and ventral temporal areas?

Response #11: We thank the reviewer for pointing out this issue. We added a statement about the ventral visual cortex in describing the location of the ROI and added the ventral view of ROIs in the Figure 1. The language-responsive and math -responsive ROIs covers both the lateral and ventral visual cortex, whereas executive function (response-conflict) regions cover only the lateral visual cortex. We compared the connectivity patterns of these three regions and found no differences (see supplementary Fig S2).

The blind group results are characterized as reflecting a reorganization in comparison to sighted adults while the results for sighted adults compared to infants are discussed more as a maturation ("adult pattern isn't default but requires experience to establish"). Both the sighted and blind adult groups showed differences from the infant group, and these differences are attributed to the role of experience. Why use "reorganization" for one result and maturation for another?

Response #12: We agree with the reviewer that both of the adult groups should be thought of as equal in relation to the infants. In other words, the brain develops under one set of experiential conditions or another. We do not think that the adult sighted pattern reflects maturation. Rather, the sighted adult pattern reflects the combined influence of maturation and visual experience. The adult blind pattern reflects the combined influence of maturation and blindness. We use the term ‘reorganization’ to label differences in the blind adults relative to sighted infants. We do so for the purpose of clarity and to remain consistent with terminology in prior liaterature. However, we agree with the reviewer that the blind group does not reflect ‘reorganization’ intrinsically any more than the sighted adult group.

The statement that "visual experience is required to set up long-range functional connectivity" is unclear, especially since the infant and blind groups showed stronger long-range functional correlations with PFC.

Response #13: We revised this sentence to specifically as “visual experience establishes elements of the sighted-adult long-range connectivity” in tha Abstract line 17.

The statement that the visual ROIS roughly correspond to "the anatomical location of areas such as V5/MT+, LO, V3a, and V4v" appears imprecise. From Supplementary Figure S7, these areas cover anterior portions of ventral temporal cortex (do these span the anatomical location of putative category-selective areas?) and into the intraparietal sulcus.

Response #14: Thanks to the reviewer for the clarification. The ventral ROIs cover the middle and part of the anterior portion of the ventral temporal lobe, including the putative category-selective areas. Additionally, the dorsal ROIs extend beyond the occipital lobe to the intraparietal sulcus and superior parietal lobule. We have added a more detailed description of the anatomical location of the ROI in the Methods section Page 17 line 489 as follows:

“Each functional ROI spans multiple anatomical regions and together the secondary visual ROIs tile large portions of lateral occipital, occipito-temporal, dorsal occipital and occipito-parietal cortices. In sighted people, the secondary visual occipital ROIs include the anatomical locations of functional regions such as motion area V5/MT+, the lateral occipital complex (LO), category specific ventral occipitotemporal cortices and dorsally, V3a and V4v. The occipital ROI also covers the middle of the ventral temporal lobe. Dorsally, it extended to the intraparietal sulcus and superior parietal lobule.”

The motivation for assessing correlations with motor and frontal regions was briefly discussed in the introduction. It would be helpful to reiterate this motivation when first introducing the analyses in the results.

Response #15: Thank you for the thoughtful suggestion. Upon reflection, we chose to substantially revise the Introduction to more clearly and comprehensively explain the rationale for examining the couplings with motor and frontal regions, rather than reiterating it in the Results section. We believe this revised framing provides a stronger foundation for the analyses that follow, while avoiding redundancy across sections. We hope this addresses the reviewer’s concern.

Reviewer #2 (Recommendations for the authors):

Congratulations on a well-written paper and an interesting set of results.

Reviewer #3 (Recommendations for the authors):

Abstract:

Mentioning "sighted infants" does not seem adequate.

Response #16: In our dataset, newborns (average age at scan = 2.79 weeks) have very limited and immature vision. We agree with the reviewer that long-term visual outcomes cannot be guaranteed without follow-up data. The term "sighted infants" was used operationally to distinguish this cohort from congenitally blind populations.

In sentences after "Specifically...", it was not clear whether the authors referred to V1 connectivity.

Response #17: We thank the reviewer for this comment. In the revised abstract, we have removed the original "Specifically..." phrasing and clarified the results.

Introduction

Talking about the "instructive effects" of vision might be confusing or misleading. Visual experiences like exposure to oral language are part of the normal/spontaneous environment that allows the infant behavioral acquisitions (contrarily with learnings that occur later during development with instruction like for reading).

Response #18: We appreciate the reviewer’s concern and would like to clarify that the term “instructive effect” is used here derived from neurodevelopmental studies (Crair, 1999; Sur et al., 1999). In this context, “instructive” refers to activity-dependent mechanisms where patterns of neural activity actively guide the organization of synaptic connectivity, emphasizing that spontaneous or sensory-driven activity (e.g., retinal waves, visual experience) can directly shape circuit refinement, as seen in ocular dominance column formation. In the context of our study, we emphasize that vision plays an instructive role in setting up the balance of connectivity between occipital cortex and non-visual networks.

For references on the development of connectivity, I would advise citing MRI studies but also studies based on histological approaches (see for example the detailed review by Kostovic et al, NeuroImage 2019).

Response #19: We thank the reviewer for this suggestion. We have incorporated a discussion on the long-range anatomical connections that emerge as early as infancy, referencing studies that employed diffusion MR imaging and histological methods, as detailed below.

“Many long-range anatomical connections between brain regions are already established in infants, even before birth, although they are not yet mature (Huang et al., 2009; Kostović et al., 2019, 2021; Takahashi et al., 2012; Vasung, 2017).” (Page 12, line 303 in the manuscript)

Results

P7 l170: It might be helpful to be precise that this is "compared with inter-hemispheric connectivity".

Response #20: We thank the reviewer for this suggestion. To align with our established terminology, we have revised the statement to explicitly contrast within-hemisphere connectivity with between-hemisphere connectivity. The modified text now reads (page 7, line 183 in the manuscript):

“Compared to sighted adults, blind adults exhibited a stronger dominance of within-hemisphere connectivity over between-hemisphere connectivity. That is, in people born blind, left visual networks are more strongly connected to left PFC, whereas right visual networks are more strongly connected to right PFC.

L176-181: It was not clear to me what was the difference between "across" and "between hemisphere connectivity". Would it be informative to test the difference between blind and sighted adults?

Response #21: We clarify that there is no distinction between the terms “across” and “between hemisphere connectivity”—they refer to the same concept. To ensure consistency, we have revised the text to exclusively use “between hemisphere connectivity” throughout the manuscript. Regarding the comparison between blind and sighted adults, we conducted statistical comparisons between these groups in our analysis, and the results have been incorporated into the revised version (Page 7, line 187 in the manuscript).

Adding statistics on Figure 3, but also on Figures 1 and 2 might help the reading.

Response #22: We have added the statistics in Figure 1-4.

Adding the third comparison in Figure 4 would be possible in my view.

Response #23: We explored integrating the response-conflict region into Figure 4, but this would require a 3x3 bar chart with pairwise statistical significance markers, which introduced excessive visual complexity that hindered readers’ ability to grasp our intended message. To ensure clarity, we retained the original Figure 4 while providing the complete three-region analysis (including all statistical comparisons) in Supplementary Figure S8 to ensure completeness.

Methods

The authors might have to specify ages at birth, and ages at scan (median + range?).

Response #24: We have added that information in the Methods section as follows:

“The average age from birth at scan = 2.79 weeks (SD = 3.77, median = 1.57, range = 0 – 19.71); average gestational age at scan = 41.23 weeks (SD = 1.77, median = 41.29, range = 37 – 45.14); average gestational age at birth = 38.43 weeks (SD = 3.73, median = 39.71, range = 23 – 42.71).” (Page 14, line 379 in the manuscript)

It might be relevant to comment on the range of available fMRI volumes, and the fact that connectivity measures might then be less robust in infants.

Response #25: We report the range of fMRI volumes in the Methods section (Page 16, Line 449). Adult participants (blind and sighted) underwent 1–4 scanning sessions, each containing 240 volumes (mean scan duration: 710.4 seconds per participant). For infants, all subjects had 2300 fMRI volumes, and we retained a subset of 1600 continuous volumes per subject with the minimum number of motion outliers. While infant connectivity measures may inherently exhibit lower robustness due to developmental and motion-related factors, our infant cohort’s large sample size (n=475) and stringent motion censoring criteria enhance the reliability of group-level inferences. We have integrated this clarification into the Methods section (Page 16, Line 444) as follows:

"While infant connectivity estimates may be less robust at the individual level compared to adults due to shorter scan durations and higher motion, our cohort’s large sample size (n=475) and rigorous motion censoring mitigate these limitations for group-level analyses. "

The mention of dHCP 2nd release should be removed from the paragraph on data availability.

Response #26: We have removed it.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Comparisons across adults with different sensory histories (blind vs. sighted) have uncovered effects of experience on human brain function. In people born blind visual cortices are responsive to non-visual tasks and show altered functional connectivity at rest. Since almost all research has been done with adults, little is known about the developmental origins of this plasticity. Are infant visual cortices initially functionally like those of sighted adults and blindness causes reorganization? Alternatively, do infants start like blind adults, with vision required to set up the sighted pattern? To distinguish between these possibilities, we compare resting state functional connectivity across blind (n = 30) and blindfolded sighted (n = 50) adults to a large cohort of sighted infants (Developing Human Connectome Project, n = 475). Remarkably, we find that infant secondary visual cortices functionally resemble those of blind more than sighted adults, consistent with the idea that visual experience is required to set up long-range functional connectivity. Primary visual cortices show a mixture of instructive effects of vision and reorganizing effects of blindness. Specifically, in sighted adults, visual cortices show stronger functional coupling with nonvisual sensory-motor networks (i.e., auditory, somatosensory/motor) than with higher-cognitive prefrontal cortices (PFC). In blind adults, visual cortices show stronger coupling with PFC. In infants, connectivity of secondary visual cortices is stronger with PFC, while V1 shows equal sensory-motor/PFC connectivity. In contrast, lateralization of occipital-to-frontal connectivity resembles the sighted adults at birth and is reorganized by blindness, possibly due to recruitment of occipital networks for lateralized cognitive functions, such as language.

Revised version: Comparisons of visual cortex function across blind and sighted adults reveals effects of experience on human brain function. Since almost all research has been done with adults, little is known about the developmental origins of plasticity. We compared resting state functional connectivity of visual cortices of blind adults (n = 30), blindfolded sighted adults (n = 50) to a large cohort of sighted infants (Developing Human Connectome Project, n = 475). Visual cortices of sighted adults show stronger coupling with non-visual sensory-motor networks (auditory, somatosensory/motor), than with higher-cognitive prefrontal cortices (PFC). In contrast, visual cortices of blind adults show stronger coupling with higher-cognitive PFC than with nonvisual sensory-motor networks. Are infant visual cortices functionally like those of sighted adults? Alternatively, do infants start like blind adults, with vision required to set up the sighted adult pattern? Remarkably, we find that, in infants, secondary visual cortices are more like those of blind adults: stronger coupling with PFC than with nonvisual sensory-motor networks, suggesting that visual experience establishes elements of the sighted-adult long-range connectivity. Infant primary visual cortices are in-between blind and sighted adults i.e., equal PFC and sensory-motor connectivity. The lateralization of occipital-to-frontal connectivity in infants resembles the sighted adults, consistent with reorganization by blindness. These results reveal instructive effects of vision and reorganizing effects of blindness on functional connectivity.

Introduction

Earlier version: Studies of vision loss provide insight into how early life experience shapes cortical function and behavior. Visual cortices of adults born blind show enhanced responses during non-visual tasks, such as Braille reading by touch, spoken language comprehension and auditory spatial processing (Bedny et al., 2011; Burton et al., 2012; Collignon et al., 2011; Kanjlia et al., 2016, 2021; Lane et al., 2015; Masuda et al., 2021; N. Raz et al., 2005; Sadato et al., 1996). Occipital cortices of blind adults also show distinctive patterns of spontaneous neural activity and functional connectivity (resting state correlations) with non-visual networks (Abboud & Cohen, 2019; Bedny et al., 2011; Burton et al., 2014; Butt et al., 2013; Deen et al., 2015; Y. Liu et al., 2007; Striem-Amit et al., 2015; Watkins et al., 2012). The developmental origins of these function differences across blind and sighted people are not known, since almost all research is done with adults.

Revised version: Relative to sighted adults, visual cortices of adults born blind show enhanced responses during non-visual tasks, such as reading braille and localizing sounds as well as distinctive patterns of long-range functional connectivity with non-visual networks (Abboud & Cohen, 2019; Bedny et al., 2011; Burton et al., 2012, 2014; Butt et al., 2013; Collignon et al., 2011; Deen et al., 2015; Kanjlia et al., 2016, 2021; Lane et al., 2015; Liu et al., 2007; Sadato et al., 1996; Striem-Amit et al., 2015; Watkins et al., 2012). Since almost all research thus far has been done with adults, a key outstanding question concerns the developmental origins of these experience-based differences.

Earlier version: One possibility is that at birth infants start out in the ‘prepared’ sighted adult state and differences among sighted and blind adults reflects reorganization caused by blindness. However, it is also possible that some of the functional differences observed between blind and sighted adults reflect the lack of instructive effects of visual experience. In other words, infants start out similar to blind adults and vision ‘instructs’ the sighted adult pattern. To distinguish between these possibilities, we compare the functional connectivity profile of visual cortices of blind and sighted adults to a large cohort of sighted infants on average 2-weeks old (Developing Human Connectome Project, dHCP, n = 475). We used resting state data as a common measure of cortical function across these diverse populations.

Revised version: One possibility is that at birth, infant visual cortices start out in the ‘prepared’ sighted adult state and blindness leads to functional reorganization. Alternatively, infant visual cortices may start out functionally similar to those of blind adults and lifetime visual experience establishes the sighted adult pattern. To distinguish between these possibilities, we compare the function of visual cortices across blind adults, sighted adults and a large cohort of 2-week-old sighted infants (Developing Human Connectome Project, dHCP, n = 475). We examined functional connectivity using resting state data, which provide a common measure of cortical function across these diverse populations.

Earlier version: Previous resting state studies comparing sighted infants to sighted adults have largely reported similarities, implying a ‘prepared’ sighted pattern (Barttfeld et al., 2018; Doria et al., 2010; Fransson et al., 2009; Gao et al., 2009; W. C. Liu et al., 2008; Zhang et al., 2019). However, these prior studies focused mostly on connectivity within large scale functional networks (e.g., visual areas are more correlated with other visual areas than with somatosensory networks), whereas differences between blind and sighted adults are observed in connectivity between the visual system and non-visual functional networks i.e., which non-visual networks are most correlated with the visual system differs across blind an sighted adults.

Revised version: To our knowledge no prior studies have compared infants to multiple populations of adults with different sensory experiences. Previous studies comparing sighted infants to sighted adults have largely reported similarity across groups (Barttfeld et al., 2018; Doria et al., 2010; Fransson et al., 2009; Gao et al., 2009; Liu et al., 2008; Zhang et al., 2019). However, these studies focused on whether large scale functional networks are present in infancy e.g., stronger connectivity of regions within the visual network, than between visual and auditory regions. Studies comparing blind and sighted adults find differences across groups in which non-visual networks are most strongly coupled with the visual system i.e., visual cortices of sighted adults show stronger coupling with non-visual sensory-motor networks (i.e., auditory, somatosensory/motor) than higher-cognitive systems; by contrast, in blind adults, visual cortex coupling is stronger with higher-cognitive prefrontal cortices (PFC) than with non-visual sensory-motor networks (Abboud & Cohen, 2019; Bedny et al., 2011; Burton et al., 2014; Deen et al., 2015; Kanjlia et al., 2021; Liu et al., 2007; Qin et al., 2013; Striem-Amit et al., 2015; Watkins et al., 2012; Yu et al., 2008). In the current study, we compare this experience-sensitive functional signature across infants, sighted and blind adults.

Added in the revised version: We measured the connectivity profile of four occipital ‘visual’ areas that show cross-modal plasticity in blindness, i.e., respond to non-visual tasks in blind people: Three secondary visual areas (located in lateral, dorsal and parts of the ventral occipital cortex) and the primary visual cortex (V1).

Earlier version: We examined the connectivity profile of three secondary visual areas on the lateral, dorsal and ventral occipital surface that have previously been found to respond to different non-visual tasks in blind adults and show different functional connectivity across blind and sighted adults (Kanjlia et al., 2016, 2021; Lane et al., 2015). In blind adults, these three regions respond to different higher-cognitive tasks, including language, numerical reasoning and executive control (Kanjlia et al., 2016, 2021; Lane et al., 2015). In sighted people, these regions roughly correspond to the anatomical location of areas such as motion area V5/MT+, the lateral occipital complex (LO), V3a and V4v (Tootell et al., 1997; Van Essen et al., 2001). Although their precise visual function in sighted people is not known. Together, these three regions tile the majority of lateral occipital cortex, providing a good sample of the connectivity profile of higher-order visual areas.

Revised version: The three secondary visual areas have been found to respond to different non-visual tasks in blind people: language tasks, numerical reasoning tasks and executive control tasks respectively. Enhanced coupling with PFC in blindness is observed across all of three occipital regions. However, in blind people each region shows preferential coupling with distinct subregions of PFC with analogous functional profiles (e.g., language responsive occipital areas are more coupled with language responsive PFC) (Kanjlia et al., 2016, 2021; Lane et al., 2015). The precise visual functions of these secondary visual regions in sighted people are not known. Anatomically, these regions correspond roughly to the location of areas such as motion area V5/MT+, the lateral occipital complex (LO), V3a and V4v in sighted people (Tootell et al., 1997; Van Essen et al., 2001).

Earlier version: We also examined connectivity of primary visual cortex (V1), which likewise shows altered task-based responses and functional connectivity in congenitally blind adults (Amedi et al., 2003; Bedny et al., 2011; Burton et al., 2014; Butt et al., 2013; Lane et al., 2015; N. Raz et al., 2005; Sadato et al., 1996; Striem-Amit et al., 2015; Yu et al., 2008). Since many previous studies have found that blindness alters the balance of connectivity between visual cortex and higher-order prefrontal as opposed to sensory-motor regions, this was our primary outcome measure. We also examined changes in connectivity lateralization – i.e., the balance of between vs. within hemisphere connectivity.

Revised version: We also examined connectivity of anatomically defined primary visual cortex (V1), which likewise shows altered task-based responses and functional connectivity in congenitally blind adults (Amedi et al., 2003; Bedny et al., 2011; Burton et al., 2014; Butt et al., 2013; Lane et al., 2015; Raz et al., 2005; Sadato et al., 1996; Striem-Amit et al., 2015; Yu et al., 2008). Since many previous studies have found that blindness alters the balance of connectivity between visual cortex and higher-order prefrontal as opposed to sensory-motor regions, this was our primary outcome measure. We also examined changes in connectivity lateralization – i.e., the balance of between vs. within hemisphere connectivity.

Earlier version: To preview the results, we find that the functional connectivity of secondary visual areas in infants resembles that of blind more than sighted adults, whereas V1 of infants falls between the two adult populations. This suggests that vision plays an instructive role in setting up the balance of connectivity between visual cortex and non-visual networks. In contrast, lateralization patterns appear to reflect blindness-related reorganization.

Revised version: To preview the results, we find that, in infants, the long-range functional connectivity profile of secondary visual areas resembles that of blind adults, whereas V1 falls between blind and sighted adult populations. Relative to sighted adults, both blind adults and infants show stronger coupling between visual cortices and PFC and weaker coupling between visual cortices and non-visual sensory-motor networks. This suggests that vision plays an instructive role in setting up the balance of connectivity between occipital cortex and non-visual networks. In contrast, connectivity lateralization patterns appear to reflect blindness-related reorganization.

Results

Connectivity profile of secondary visual cortices in sighted infants is more similar to that of blind than sighted adults

Earlier version: Secondary visual areas of sighted adults showed stronger functional connectivity with non-visual sensory areas (primary somatosensory and motor cortex, S1/M1, and primary auditory cortex, A1) than with prefrontal cortices (PFC). By contrast, in blind adults, visual cortices showed higher functional connectivity with PFC than with non-visual sensory areas (S1/M1 and A1) (group (sighted adults, blind adults) by ROI (PFC, non-visual sensory) interaction effect: F(1, 78) = 148.819, p < 0.001; post-hoc Bonferroni-corrected paired t-test, sighted adults: non-visual sensory > PFC: t (49) = 9.722, p < 0.001; blind adults: non-visual sensory < PFC: t(29) =8.852, p < 0.001; Fig. 1).

Revised version: We first examined the long-range functional connectivity of (three) secondary visual areas with sensory-motor areas on the one hand, and higher-order PFC networks on the other. In sighted adults, all three secondary visual areas showed stronger functional connectivity with non-visual sensory-motor areas (primary somatosensory and motor cortex, S1/M1, and primary auditory cortex, A1) than with higher-cognitive prefrontal cortices (PFC). By contrast, in blind adults, all secondary visual regions showed stronger functional connectivity with PFC than with non-visual sensory-motor areas (S1/M1 and A1) (group (sighted adults, blind adults) by ROI (PFC, non-visual sensory) interaction effect: F(1, 78) = 148.819, p < 0.001; post-hoc Bonferroni-corrected paired t-test, sighted adults: non-visual sensory > PFC: t (49) = 9.722, p < 0.001; blind adults: non-visual sensory < PFC: t (29) =8.852, p < 0.001; Fig. 1).

Functional connectivity of secondary visual cortices.

Earlier version: (A) Bar graph shows functional connectivity (r) of secondary visual cortices (blue) to non-visual sensory motor areas (purple) and prefrontal cortices (green), averaged across occipital, PFC and sensory-motor ROIs (A1 and S1/M1) in sighted adults, blind adults and sighted infants. Regions of interest (ROI) displayed on the left. Note that regions extend to ventral surface, not shown. See Supplementary Figure S7 for the full views of ROIs. (B) Circle plots represent the connectivity of secondary visual cortices to non-visual networks, min-max normalized to [0,1], i.e., as a proportion. OC: occipital cortices; MTH: math-responsive region; LG: language-responsive region; EF: executive function-responsive (response-conflict) region.

Revised version: (A) Bar graph shows functional connectivity (r) of secondary visual cortices (blue) to non-visual sensory motor areas (purple) and prefrontal cortices (green), averaged across occipital, PFC and sensory-motor ROIs (A1 and S1/M1) in sighted adults, blind adults and sighted infants. Regions of interest (ROI) displayed on the left. Note that regions extend to ventral surface, not shown. See Supplementary Figure S5 for the full views of three occipital ROIs. (B) Circle plots represent the connectivity of secondary visual cortices to non-visual networks, min-max normalized to [0,1], i.e., as a proportion. OC: occipital cortices; MTH: math-responsive region; LG: language-responsive region; EF: executive function-responsive (response-conflict) region. Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05).

Earlier version: Like in blind adults, in sighted infants, secondary visual cortices showed higher connectivity to PFC than non-visual sensory areas (S1/M1 and A1) (non-visual sensory < PFC paired t-test, t (474) = 20.144, p < 0.001) (Fig. 1). The connectivity matrix of sighted infants was also more correlated with that of blind than sighted adults, but strongly correlated with both adult groups (secondary visual, PFC and non-visual sensory areas: sighted infants correlated to blind adults: r = 0.721, p < 0.001; to sighted adults: r = 0.524, p < 0.001; difference between correlations of infants to blind vs. to sighted adults: z = 3.77, p < 0.001; see Supplementary Figure S1 for the connectivity matrices).

Revised version: Like in blind adults, in sighted infants, secondary visual areas showed higher connectivity to PFC than to non-visual sensory-motor areas (S1/M1 and A1) (non-visual sensory < PFC paired t-test, t (474) = 20.144, p < 0.001) (Fig. 1). The connectivity matrix of sighted infants was more correlated with that of blind than sighted adults, but strongly correlated with both adult groups (secondary visual, PFC and non-visual sensory areas: sighted infants correlated to blind adults: r = 0.721, p < 0.001; to sighted adults: r = 0.524, p < 0.001; difference between correlations of infants to blind vs. to sighted adults: z = 3.77, p < 0.001; see Supplementary Figure S1 for the connectivity matrices).

Earlier version: These results suggests that vision is required to set up the sighted adult functional connectivity pattern, i.e., vision enhances occipital cortex connectivity to non-visual sensory-motor networks and dampens connectivity to higher-cognitive prefrontal networks.

Revised version: These results suggests that vision is required to set up elements of the sighted adult functional connectivity pattern, i.e., vision enhances occipital cortex connectivity to non-visual sensory-motor networks and dampens connectivity to higher-cognitive prefrontal networks.

Earlier version: We checked the robustness of these results in a number of ways. First, we compared the effects across the three different secondary visual regions and found that the same pattern held across all three regions (Supplementary results and Supplementary Figure S2). Next, to check the robustness of the findings in infants we randomly split the infant dataset into two halves and did split-half cross-validation. Across all comparisons the results of the two halves were highly similar, suggesting the effects are robust (see Supplementary Figure S3 to Supplementary Figure S6). We performed this validation procedure for all analyses reported below with similar results.

Revised version: We checked the robustness of these results in a number of ways. We first compared the effects across the three secondary visual regions and observed the same pattern across all. (Supplementary results and Supplementary Figure S2). Next, to check the robustness of the findings in infants we randomly split the infant dataset into two halves and did split-half cross-validation. Across all comparisons the results of the two halves were highly similar, suggesting the effects are robust (see Supplementary Figure S3 and S4). We performed this validation procedure for all analyses reported below with similar results.

Earlier version: The connectivity pattern of V1 influenced both by early visual experience and blindness

Revised version: The connectivity pattern of V1 is influenced by early visual experience and blindness

Earlier version: As for secondary visual corticies, we examined the functional connectivity of the primary visual cortex (V1) with non-visual sensory areas (S1/M1 and A1) and PFC. V1 showed the same dissociation between sighted and blind adults as secondary visual areas: in sighted adults, V1 has stronger functional connectivity with non-visual sensory areas than with PFC. By contrast, in blind adults, V1 shows stronger connectivity with PFC than with non-visual sensory areas (group (sighted adults, blind adults) by ROI (PFC, non-visual sensory) interaction: F(1, 78) = 125.775, p < 0.001; post-hoc Bonferroni-corrected paired t-test, sighted adults non-visual sensory > PFC: t (49) = 9.404, p < 0.001; blind adults non-visual sensory < PFC: t (29) =7.128, p < 0.001; Fig. 2).

Revised version: V1 showed the same dissociation between sighted and blind adults as secondary visual areas: in sighted adults, V1 has stronger functional connectivity with non-visual sensory-motor areas than with PFC. By contrast, in blind adults, V1 shows stronger connectivity with PFC than with non-visual sensory areas (group (sighted adults, blind adults) by ROI (PFC, non-visual sensory) interaction: F(1, 78) = 125.775, p < 0.001; post-hoc Bonferroni-corrected paired t-test, sighted adults non-visual sensory > PFC: t (49) = 9.404, p < 0.001; blind adults non-visual sensory < PFC: t (29) =7.128, p < 0.001; Fig. 2).

Functional connectivity of primary visual cortices (V1).

Earlier version: Regions of interest (ROI) displayed on the upper. Bar graph shows functional connectivity (r) of V1 to non-visual sensory motor areas (purple) and prefrontal cortices (green), averaged across three PFC ROIs and sensory-motor ROIs (S1/M1 and A1).

Revised version: Regions of interest (ROI) displayed on the upper. Bar graph shows functional connectivity (r) of V1 to non-visual sensory motor areas (purple) and prefrontal cortices (green), averaged across three PFC ROIs and sensory-motor ROIs (S1/M1 and A1). Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05). Cross (†) denote marginal difference in Bonferroni-corrected pairwise comparisons (0.05< p <0.1).

The pattern for sighted infants in V1 fell between that of sighted and blind adults. The connectivity matrix of sighted infants (V1, PFC, and non-visual sensory) was equally correlated with blind and sighted adults (infants correlated to blind adults: r = 0.654, p < 0.001; to sighted adults: r = 0.594, p < 0.001; correlation of infants with blind vs. with sighted adults: z = 0.832, p = 0.406; see Supplementary Figure S1 for the connectivity matrices). The difference in connectivity strength between V1 to PFC and V1 to non-visual sensory regions was weaker in sighted infants than in sighted or blind adults (group (sighted adults, infants) by ROI (PFC, non-visual sensory) interaction effect: F(1, 523) = 92.21, p < 0.001; group (blind adults, infants) by ROI (PFC, non-visual sensory) interaction effect: F(1, 503) = 57.444, p < 0.001). V1 of sighted infants showed marginally stronger connectivity to non-visual sensory regions (S1/M1 and A1) than PFC (non-visual sensory regions > PFC, paired t-test, t (474) = 1.95, p = 0.052; Fig.2).

Earlier version: The dHCP cohort included both full-term neonates and preterm infants, scanned at their equivalent gestational age. Visual exposure therefore varied somewhat in duration across infants (from 0 to 19.71 weeks), with slightly longer exposure in preterm babies. This variation did not affect connectivity patterns either in V1 or secondary visual cortices (V1: r = 0.06, p = 0.192; secondary visual: r = 0.004, p = 0.923; see Supplementary Figure S8). A few weeks of vision after birth is therefore insufficient to influence connectivity.

Revised version: The dHCP cohort included both full-term neonates and preterm infants, scanned at their equivalent gestational age. Visual exposure therefore varied somewhat in duration across infants (from 0 to 19.71 weeks), with slightly longer exposure in preterm babies. This variation did not affect connectivity patterns either in V1 or secondary visual cortices (V1: r = 0.06, p = 0.192; secondary visual: r = 0.004, p = 0.923; see Supplementary Figure S6). We also compared the connectivity patterns of preterm (n = 90) and full-term infants (n = 385) and found no difference from each other or from the all-infant dataset (see Supplementary Figure S7). A few weeks of vision after birth is therefore insufficient to influence connectivity.

Evidence for blindness-related reorganization in laterality of occipito-frontal connectivity

Earlier version: Relative to sighted adults, blind adults showed a stronger dominance of within hemisphere connectivity between secondary and primary visual cortices and PFC. That is, in people born blind, left visual networks are more strongly connected to left PFC networks, whereas right visual networks are more strongly connected to right PFC. By contrast, there is no hemispheric bias in the sighted group (group (blind adults, sighted adults) by lateralization (within hemisphere, between hemisphere) interaction, secondary visual cortices: F(1, 78) = 131.51, p < 0.001; post-hoc Bonferroni-corrected paired: t-test: sighted adults within hemisphere > across hemisphere: t (49) = 5.778, p < 0.001; blind adults within hemisphere > across hemisphere: t(29) = 10.735, p < 0.001; V1: F(1, 78) = 87.211, p < 0.001; post-hoc Bonferroni-corrected paired: t-test: sighted adults within hemisphere > between hemisphere: t (49) = 3.251, p = 0.101; blind adults within hemisphere > between hemisphere: t (29) = 7.019, p < 0.001).

Revised version: Compared to sighted adults, blind adults exhibit a stronger dominance of within-hemisphere connectivity over between-hemisphere connectivity. That is, in people born blind, left visual networks are more strongly connected to left PFC, whereas right visual networks are more strongly connected to right PFC. By contrast, in sighted adults, this cross-hemisphere difference is weak or absent. This difference between adult groups is observed for both V1 and secondary visual cortices (group (blind adults, sighted adults) by lateralization (within hemisphere, between hemisphere) interaction in secondary visual cortices: F(1, 78) = 131.51, p < 0.001). Secondary visual cortices showed a significant within > between difference in both groups, with a lager effect in the blind group (post-hoc tests, Bonferroni-corrected paired: t-test: sighted adults within hemisphere > between hemisphere: t (49) = 7.441, p = 0.012; blind adults within hemisphere > between hemisphere: t (29) = 10.735, p < 0.001; V1: F(1, 78) = 87.211, p < 0.001). In V1, only the blind group showed a significant within > between hemisphere effect (post-hoc Bonferroni-corrected paired: t-test: sighted adults within hemisphere < between hemisphere: t (49) = 3.251, p = 0.101; blind adults within hemisphere > between hemisphere: t (29) = 7.019, p < 0.001).

Removed in the revised version: These connectivity results are consistent with previous studies of task-based cross-modal responses in blindness. In blind adults cross-modal responses in occipital cortex and co-lateralize with fronto-parietal networks with related functions (Kanjlia et al., 2021; Lane et al., 2017). For example, language-responsive occipital areas collateralize with language responsive prefrontal areas across individuals (Lane et al., 2017).

Earlier version: The connectivity pattern of sighted infants resembled sighted more than blind adults (Fig. 3), suggesting blindness-driven reorganization. There was a significant difference in laterality between blind adults and sighted infants (group (blind adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 503) = 303.04, p < 0.001). By contrast, there was no difference between sighted adults and sighted infants (group (sighted adults, infants) by lateralization (within hemisphere, across hemisphere) interaction effect: F(1, 523) = 2.244, p = 0.135; see supplementary results for a detailed group comparison of within and across hemisphere differences). Similar group by laterality interaction pattern are also observed in V1 (group (blind adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 503) = 123.608, p < 0.001; group (sighted adults, infants) by lateralization (within hemisphere, across hemisphere) interaction effect: F(1, 523) = 2.827, p = 0.093). The incorporation of visual cortices into lateralized functional networks (e.g., language, response selection) in blindness, which is observed in task-based studies, may drive stronger within-hemisphere connectivity in this population (Kanjlia et al., 2021; Lane et al., 2017; Tian et al., 2022).

Revised version: With respect to laterality, infants resembled sighted more than blind adults (Fig. 3). For secondary visual cortices, there was a significant difference between blind adults and sighted infants and no difference between sighted adults and sighted infants (group (blind adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 503) = 303.04, p < 0.001); (group (sighted adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 523) = 2.244, p = 0.135). A similar group by laterality interaction were observed for V1 (group (blind adults, infants) by lateralization (within hemisphere, between hemisphere) interaction: F(1, 503) = 123.608, p < 0.001; group (sighted adults, infants) by lateralization (within hemisphere, between hemisphere) interaction effect: F(1, 523) = 2.827, p = 0.093). This suggests that the enhancement of within over between hemisphere long-range connectivity is related to blindness-driven reorganization.

Earlier version: Within hemisphere vs. across hemisphere functional connectivity.

Revised version: Within hemisphere vs. between hemisphere functional connectivity.

Earlier version: Bar graph shows within hemisphere (blue) and across hemisphere (orange) functional connectivity (r coefficient of resting state correlations) of secondary visual (left) and V1 (right) to prefrontal cortices in sighted adults, blind adults, and sighted infants. Blind adults show a larger difference than any of the other groups.

Revised version: Bar graph shows within hemisphere (blue) and between hemisphere (orange) functional connectivity (r coefficient of resting state correlations) of secondary visual (left) and V1 (right) to prefrontal cortices in sighted adults, blind adults, and sighted infants. Blind adults show a larger difference than any of the other groups. Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05).

Added in the revised version: Task-based fMRI studies find that cross-modal responses in occipital cortex colateralize with fronto-parietal networks with related functions (e.g., language, response selection) (Kanjlia et al., 2021; Lane et al., 2017). For example, language-responsive occipital areas collateralize with language responsive prefrontal areas across individuals (Lane et al., 2017). Recruitment of visual cortices by cross-modal tasks (e.g., spoken language) may enhance within-hemisphere connectivity in people born blind (Kanjlia et al., 2021; Lane et al., 2017; Tian et al., 2022). Together, this evidence supports the hypothesis that, starting from the less lateralized infant state, blindness increases lateralization of occipital long-range connectivity.

Earlier version: Specialization across different fronto-occipital networks: present in adults, absent at birth

Revised version: Specialization of connectivity across different fronto-occipital networks: present in adults, absent at birth

Earlier version: It is not the case that every occipital area shows equal connectivity to every prefrontal area. On the contrary, in blind adults resting state connectivity patterns are specialized and aligned with the functional specialization observed in task-based data (Bedny et al., 2011; Kanjlia et al., 2016, 2021). For example, language-responsive subregions of occipital cortex show strongest functional connectivity with language-responsive sub-regions of PFC, whereas math-responsive occipital areas show stronger connectivity with math-responsive PFC (Fig. 4) (Bedny et al., 2011; Kanjlia et al., 2016, 2021; Lane et al., 2015). Is this fronto-occipital connectivity specialization present in infancy, potentially driving task-based cross-modal specialization?

Revised version: In blind adults, different occipital areas show enhanced connectivity patterns with distinct subregions of PFC and this specialization is aligned with the functional specialization observed in task-based data (Bedny et al., 2011; Kanjlia et al., 2016, 2021). For example, language-responsive subregions of occipital cortex show strongest functional connectivity with language-responsive sub-regions of PFC, whereas math-responsive occipital areas show stronger connectivity with math-responsive PFC. This pattern is most pronounced in blind people but can be seen weakly even in sighted participants (Fig. 4) (Bedny et al., 2011; Kanjlia et al., 2016, 2021; Lane et al., 2015). Is this fronto-occipital connectivity specialization present in infancy, potentially enabling the task-based cross-modal specialization?

Occipito-frontal functional connectivity.

Earlier version: Bar graph shows across functional connectivity of different sub-regions of prefrontal (PFC) and occipital cortex (OCC) in sighted adults, blind adults, and sighted infants. Sub-regions (regions of interest) were defined based on task-based responses in a separate dataset of sighted (frontal) and blind (frontal and occipital) adults (Kanjlia et al., 2016, 2021; Lane et al., 2015). PFC/OCC-MATH: math-responsive regions were more active when solving math equations than comprehending sentences. PFC/OC-LANG: language-responsive regions were more active when comprehending sentences than solving math equations (Kanjlia et al., 2016, 2021; Lane et al., 2015). In blind adults these regions show biases in connectivity related to their function i.e., language-responsive PFC is more correlated with language responsive OCC. No such pattern is observed in infants. See Supplementary Figure S10 for connectivity matrix.

Revised version: Bar graph shows across functional connectivity of different sub-regions of prefrontal (PFC) and occipital cortex (OCC) in sighted adults, blind adults, and sighted infants. Sub-regions (regions of interest) were defined based on task-based responses in a separate dataset of sighted (frontal) and blind (frontal and occipital) adults (Kanjlia et al., 2016, 2021; Lane et al., 2015). PFC/OCC-MATH: math-responsive regions were more active when solving math equations than comprehending sentences. PFC/OC-LANG: language-responsive regions were more active when comprehending sentences than solving math equations (Kanjlia et al., 2016, 2021; Lane et al., 2015). In blind adults these regions show biases in connectivity related to their function i.e., language-responsive PFC is more correlated with language responsive OCC. No such pattern is observed in infants. Asterisks (*) denote significant Bonferroni-corrected pairwise comparisons (p < .05). See Supplementary Figure S9 for connectivity matrix.

Earlier version: We compared connectivity preferences across three prefrontal and three occipital regions previous shown to activate preferentially in language (sentences > math), math (math > sentences) and response-conflict (no-go > go with tones) tasks respectively (Kanjlia et al., 2016, 2021; Lane et al., 2015). For ease of viewing, Fig. 4 shows results from two of the three regions, math and language. However, all statistical analyses included all three areas (See Supplementary Figure S9 for all three regions).

Revised version: We compared connectivity preferences across three prefrontal and three occipital regions previous shown to activate preferentially in language (sentences > math), math (math > sentences) and response-conflict (no-go > go with tones) (Kanjlia et al., 2016, 2021; Lane et al., 2015). For ease of viewing, Fig. 4 shows results from two of the three regions, math and language. See Supplementary Figure S8 for all three regions. Note that the statistical analyses included all three areas.

Earlier version: Sighted infants showed a less differentiated fronto-occipital connectivity pattern relative to sighted and blind adults (Group (sighted adults, blind adults, sighted infants) by occipital regions (math, language, response-conflict) by PFC regions (math, language, response-conflict) interaction F(8, 2208) = 16.323, p < 0.001). Unlike in adults, in infants, all the occipital regions showed stronger correlations with math- and response-conflict related prefrontal areas than language-responsive prefrontal areas (Fig. 4 and Supplementary Figure S9). The occipital region that is sensitive to response-conflict in blind adults showed equivalent correlations with math and response-conflict PFC regions in infants. The region of occipital cortex that responds to language and shows the strongest connectivity with language responsive PFC in blind adults, showed stronger connectivity with math and response-conflict PFC areas in infants.

Revised version: Contrary to the hypothesis that specialization of functional connectivity across different prefrontal/occipital areas is present from birth, sighted infants showed a less differentiated fronto-occipital connectivity pattern relative to both blind and sighted adults (Group (sighted adults, blind adults, sighted infants) by occipital regions (math, language, response-conflict) by PFC regions (math, language, response-conflict) interaction F(8, 2208) = 16.323, p < 0.001). Unlike in adults, in infants, all the occipital regions showed stronger correlations with math- and response-conflict related prefrontal areas than language-responsive prefrontal areas (Fig. 4 and Supplementary Figure S8). However, the preferential correlation with math responsive PFC was strongest in those occipital areas that go on to develop math responses in blind adults (occipital regions (math, language, response-conflict) by PFC regions (math, language, response-conflict) interaction in infants F(4, 1896) = 85.145, p < 0.001, post-hoc Bonferroni-corrected paired t-test see Supplementary Table 1).

Removed in the revised version: Although fronto-occipital connectivity was not adult-like in infants, biases in infants were somewhat consistent with future differentiation: the preferential correlation with math responsive PFC was strongest in those occipital areas that go on to develop math responses in blind adults (occipital regions (math, language, response-conflict) by PFC regions (math, language, response-conflict) interaction in infants F(4, 1896) = 85.145, p < 0.001, post-hoc Bonferroni-corrected paired t-test see Supplementary Table 1). Although the occipital region that is language-responsive in blind adults showed stronger connectivity to math and response-conflict areas of PFC in infants, this biased against language areas was smaller for this occipital areas than the other two.

Earlier version: Note that findings regarding regional specialization need to be interpreted with caution for two reasons. First, although, prior evidence suggests that at least some of the prefrontal specialization present in adults is already present even in young babies, we do not know whether the language/number/executive function distinctions exist in prefrontal cortices of infants (Raz & Saxe, 2020). Second, these more fine-grained comparisons across occipital/frontal regions are more vulnerable to potential anatomical alignment issues between adult and infant brains. In other words, lack of specialization in infants could reflect the different location of the areas in this population.

Revised version: Note that findings of reduced regional specialization in infants need to be interpreted with caution. First, we do not know whether the same specialization of prefrontal sub-regions seen in adults is present in infants, although, prior evidence suggests some prefrontal specialization is already present (Raz & Saxe, 2020). Second, the more fine-grained comparisons across occipital/frontal regions are more vulnerable to potential anatomical alignment issues between adult and infant brains. In other words, lack of specialization in infants could reflect the different location of the areas in this population.

Discussion

Earlier version: The sighted adult functional connectivity pattern, although the most common in the population, is not the ‘default’ starting state in infants but rather requires visual experience to establish. This was particularly evident in the case of secondary visual cortices, where, resting state connectivity patterns with non-visual networks in sighted infants resemble those of blind adults more so than those of sighted adults. Both in infants and blind adults, secondary occipital areas showed stronger functional connectivity with higher-order prefrontal cortices than with other sensory-motor networks (S1/M1, A1). Consistent with this observation, one previous study with a small sample of infants found strong connectivity between lateral occipital and prefrontal areas, although there was no comparison to blind adults in that study (Barttfeld et al., 2018). In V1, infants fell somewhere in between sighted and blind adults, suggesting an effect both of vision and of blindness on functional connectivity. We hypothesize that vision, as well as temporally coordinated multi-modal experiences contribute to establishing the sighted connectivity profile.

Revised version: The present results provide insight into the developmental process of experience-based functional specialization in human cortex. We find independent effects of visual experience and blindness on the development of visual networks. Aspects of the sighted adult connectivity pattern require visual experience to become established. This is particularly striking for secondary visual cortices, where, connectivity with nonvisual networks in sighted infants resemble blind more than sighted adults. Both in infants and blind adults, secondary occipital areas showed stronger functional connectivity with higher-order prefrontal cortices than with other sensory-motor networks (S1/M1, A1). Consistent with this observation, one previous study with a small sample of infants found strong connectivity between lateral occipital and prefrontal areas, although there was no comparison to blind adults in that study (Barttfeld et al., 2018). In V1, infants fell somewhere in between sighted and blind adults, suggesting an effect both of vision and of blindness on functional connectivity.

Added in the revised version: The present results reveal the effects of experience on development of functional connectivity between infancy and adulthood, but do not speak to the precise time course of these effects. Infants in the current sample had between 0 and 20 weeks of visual experience. Comparisons across these infants suggests that several weeks of postnatal visual experience is insufficient to produce a sighted-adult connectivity profile. The time course of development could be anywhere between a few months and years and could be tested by examining data from children of different ages.

Added in the revised version: We propose that vision, as well as temporally coordinated multi-modal experiences contribute to establishing the sighted adult connectivity profile in visual cortices. For example, coordinated visuo-motor activity during development may enhance connectivity between visual and motor networks. Supporting this idea, evidence shows that in infants, motor competence and early experience predict coupling between occipital and motor networks (Colomer et al., 2023).Despite these insights, many questions remain regarding the neurobiological mechanisms underlying experience-based functional connectivity changes and their relationship to anatomical development. Long-range anatomical connections between brain regions are already present in infants—even prenatally—though they remain immature (Huang et al., 2009; Kostović et al., 2019, 2021; Takahashi et al., 2012; Vasung, 2017). Functional connectivity changes may stem from local synaptic modifications within these stable structural pathways, consistent with findings that functional connectivity can vary independently of structural connection strength (Fotiadis et al., 2024). Moreover, functional connectivity has been shown to outperform structural connectivity in predicting individual behavioral differences, suggesting that experience-based functional changes may reflect finer-scale synaptic or network-level modulations not captured by macrostructural measures (Ooi et al., 2022). Prior studies also suggest that, even in adults, coordinated sensory-motor experience can lead to enhancement of functional connectivity across sensory-motor systems, indicating that large-scale changes in functional connectivity do not necessarily require corresponding changes in anatomical connectivity (Guerra-Carrillo et al., 2014; Li et al., 2018). Another possibility is some connectivity changes are mediated by ‘third-party’ regions, including subcortical mechanisms such as those involving the thalamus (Vega-Zuniga et al., 2025).

Added in the revised version: The current findings reveal both effects of vision and effects of blindness on the functional connectivity patterns of the visual cortex. A further open question is whether visual experience plays an instructive or permissive role in shaping neural connectivity patterns. An instructive role suggests that specific sensory experiences or patterns of neural activity directly shape and organize neural circuitry. In contrast, a permissive role implies that sensory experience or neural activity merely facilitates the influence of other factors—such as molecular signals—on the formation and organization of neural circuits (Crair, 1999; Sur et al., 1999). Studies with animals that manipulate the pattern or informational content of neural activity while keeping overall activity levels constant could distinguish between these hypotheses (Crair, 1999; Roy et al., 2020; Stellwagen & Shatz, 2002).

Earlier version: Since visual behavior is likely to be influenced by the communication of visual cortices with non-visual networks, a key question concerns the behavioral relevance of these connectivity signatures for vision and multimodal integration. An increasing number of people who grew up blind will have the possibility of sight restoration in adulthood e.g., through cataract removal, corneal transplant or gene therapy. Recent evidence suggests that sight recovery individuals show some multimodal integration deficits (Ashtari, 2020; Badde et al., 2020; Guerreiro et al., 2015; Putzar et al., 2007). There is also evidence that occipital oscillations, which affect cross-network communication, are different in this population (Pant et al., 2023). It will therefore be important to determine the contribution of connectivity to visual and multimodal sensory behavior and how connectivity might be shaped later in life.

Revised version: A further key question concerns the behavioral relevance of the connectivity signatures observed in the current study. The capacity of occipital cortices to support visual and multimodal behavior in sighted people may depend not only on local visual cortex function but also on the capacity of the visual system to coordinate its function with non-visual networks. Does enhanced connectivity between visual and non-visual sensory motor networks facilitate multimodal integration for sighted people e.g., when catching a ball? Potentially consistent with this possibility, recent evidence suggests that people who grew up blind but recover sight in adulthood show multimodal integration deficits (Ashtari, 2020; Badde et al., 2020; Guerreiro et al., 2015; Putzar et al., 2007) and distinct occipital oscillations (Pant et al., 2023).

Earlier version: For people who remain blind throughout life, the infant connectivity profile could play a role in enabling recruitment of visual cortices by non-visual functions, such as language and non-verbal executive processes. Habitual activation of occipital networks during higher cognitive tasks in early development could then itself influence both connectivity and selectivity, resulting in different adult profiles.

Revised version: Conversely, for people who remain blind throughout life, visual-PFC connectivity could enable recruitment of visual cortices for higher-order non-visual functions, such as language and executive control. Habitual activation of occipital networks during higher cognitive tasks in early development could intern enhance connectivity and specialization of visual networks for non-visual functions.

Earlier version: The clearest evidence for blindness-related reorganization in the current study was observed in the case of laterality. Connectivity lateralization in sighted infants resembles that of sighted adults, in both V1 and secondary visual cortices. Said differently, relative to both sighted infants and sighted adults, blind adults show more lateralized connectivity patterns between occipital and prefrontal cortices. Previous studies suggest that in people born blind occipital and non-occipital language responses are co-lateralized (Lane et al., 2017). We speculate that habitual activation of visual cortices by higher-cognitive tasks, such as language, which are themselves highly lateralized, contributes to this biased connectivity pattern of occipital cortex in blindness.

Revised version: In the current study, the clearest evidence for functional change driven by blindness was observed for laterality. Connectivity lateralization in sighted infants resembles that of sighted adults, in both V1 and secondary visual cortices. Relative to both sighted infants and sighted adults, blind adults show more lateralized connectivity patterns between occipital and prefrontal cortices. Previous studies suggest that in people born blind occipital and non-occipital language responses are co-lateralized (Lane et al., 2017). We speculate that habitual activation of visual cortices by higher-cognitive tasks, such as language, which are themselves highly lateralized, contributes to this biased connectivity pattern of occipital cortex in blindness.

Materials and methods

Participants

Earlier version: Fifty sighted adults and thirty congenitally blind adults contributed the resting state data (sighted: n = 50; 30 females; mean age = 35.33 years, standard deviation (SD) = 14.65; mean years of education = 17.08, SD = 3.1; blind: n = 30; 19 females; mean age = 44.23 years, SD = 16.41; mean years of education = 17.08, SD = 2.11; blind vs. sighted age, t (78) = 2.512, p < 0.05; blind vs. sighted years of education, t (78) = 0.05, p = 0.996). Since blind participants were on average older, we also performed analyses in an age-matched subgroups of sighted controls (n = 29) and found similar results to the full sample (see Supplementary Figure S11 to Figure S14). Blind and sighted participants had no known cognitive or neurological disabilities (screened through self-report). All adult anatomical images were read by a board-certified radiologist and no gross neurological abnormalities were found. All the blind participants had at most minimal light perception from birth. Blindness was caused by pathology anterior to the optic chiasm (i.e., not due to brain damage). All participants gave written informed consent under a protocol approved by the Institutional Review Board of Johns Hopkins University.

Revised version: Fifty sighted adults and thirty congenitally blind adults contributed the resting state data (sighted: n = 50; 30 females; mean age = 35.33 years, standard deviation (SD) = 14.65; mean years of education = 17.08, SD = 3.1; blind: n = 30; 19 females; mean age = 44.23 years, SD = 16.41; mean years of education = 17.08, SD = 2.11; blind vs. sighted age, t (78) = 2.512, p < 0.05; blind vs. sighted years of education, t (78) = 0.05, p = 0.996). Since blind participants were on average older, we also performed analyses in an age-matched subgroups of sighted controls (n = 29) and found similar results to the full sample (see Supplementary Figure S10). Blind and sighted participants had no known cognitive or neurological disabilities (screened through self-report). All adult anatomical images were read by a board-certified radiologist, and no gross neurological abnormalities were found. All the blind participants had at most minimal light perception from birth. Blindness was caused by pathology anterior to the optic chiasm (i.e., not due to brain damage). All participants gave written informed consent under a protocol approved by the Institutional Review Board of Johns Hopkins University.

Earlier version: Neonate data were from the third release of the Developing Human Connectome Project (dHCP) (n = 783) (https://www.developingconnectome.org). Ethical approval was obtained from the UK Health Research Authority (Research Ethics Committee reference number: 14/LO/1169). After quality control procedures (described below), 475 subjects were included in data analysis, with one scan per subject. The average age from birth at scan = 2.79 weeks (SD = 3.77); average gestational age at scan = 41.23 weeks (SD = 1.77). We only included infants who were full-term or scanned at term-equivalent age if preterm, while not being flagged by the dHCP project team as not passing quality control for functional MRI (fMRI) images (n = 634). Infants with more than 160 motion outliers were exclude (n = 116 dropped). Motion-outlier volumes were defined as DVARS (the root mean square intensity difference between successive volumes) higher than 1.5 interquartile range above the 75th centile, after motion and distortion correction. Infants with signal drop-out in regions of interest (ROI) were also excluded (n = 43 dropped). To identify signal dropout, we first averaged blood oxygen level-dependent (BOLD) signal intensity for all time point, for each subject, in each of 100 parcel defined by Schaefer’s atlas (Schaefer et al., 2018). For each ROI (n = 18 ROIs) in the current study, signal dropout was then identified as BOLD intensity lower than −3 standard deviations, where the mean and standard deviations were identified across all 100 cortical parcels. Participants were excluded if any of the ROIs showed a signal dropout.

Revised version: Neonate data were from the third release of the Developing Human Connectome Project (dHCP) (n = 783) (https://www.developingconnectome.org). Ethical approval was obtained from the UK Health Research Authority (Research Ethics Committee reference number: 14/LO/1169). After quality control procedures (described below), 475 subjects were included in data analysis, with one scan per subject. The average age from birth at scan = 2.79 weeks (SD = 3.77, median = 1.57, range = 0 – 19.71); average gestational age at scan = 41.23 weeks (SD = 1.77, median = 41.29, range = 37 – 45.14); average gestational age at birth = 38.43 weeks (SD = 3.73, median = 39.71, range = 23 – 42.71). We only included infants who were full-term or scanned at term-equivalent age if preterm, while not being flagged by the dHCP project team as not passing quality control for functional MRI (fMRI) images (n = 634). Infants with more than 160 motion outliers were exclude (n = 116 dropped). Motion-outlier volumes were defined as DVARS (the root mean square intensity difference between successive volumes) higher than 1.5 interquartile range above the 75th centile, after motion and distortion correction. Infants with signal drop-out in regions of interest (ROI) were also excluded (n = 43 dropped). To identify signal dropout, we first averaged blood oxygen level-dependent (BOLD) signal intensity for all time point, for each subject, in each of 100 parcel defined by Schaefer’s atlas (Schaefer et al., 2018). For each ROI (n = 18 ROIs) in the current study, signal dropout was then identified as BOLD intensity lower than -3 standard deviations, where the mean and standard deviations were identified across all 100 cortical parcels. Participants were excluded if any of the ROIs showed a signal dropout. The infants’ structural images were reviewed by a pediatric neuroradiologist from the dHCP team, who assigned scores on a scale from 1 to 5. A score of 1 indicated a normal appearance for the subject’s age, while scores of 4 or 5 suggested potential or likely clinical significance, or both clinical and imaging relevance. We repeated our analysis after excluding infants with a radiology score of 4 or 5, and the results remained consistent (see Figure S11).

Image acquisition

Removed in the revised version: Blind and sighted adult

Earlier version: MRI anatomical and functional images were collected on a 3T Phillips scanner at the F. M. Kirby Research Center. T1-weighted anatomical images were collected using a magnetization-prepared rapid gradient-echo (MP-RAGE) in 150 axial slices with 1 mm isotropic voxels. Resting state fMRI data were collected in 36 sequential ascending axial slices for 8 minutes. TR = 2 s, TE = 0.03 s, flip angle = 70°, voxel size = 2.4 × 2.4 × 2.5 mm, inter-slice gap = 0.5 mm, field of view (FOV) = 192 × 172.8 × 107.5. Participants completed 1 to 4 scans of 240 volume each (average scan time = 710.4 second per person). During the resting state scan, participants were instructed to relax but remain awake. Sighted participants wore light-excluding blindfolds to equalize the light conditions across the groups during the scans.

Revised version: Blind and sighted adult MRI anatomical and functional images were collected on a 3T Phillips scanner at the F. M. Kirby Research Center. T1-weighted anatomical images were collected using a magnetization-prepared rapid gradient-echo (MP-RAGE) in 150 axial slices with 1 mm isotropic voxels. Resting state fMRI data were collected in 36 sequential ascending axial slices for 8 minutes. TR = 2 s, TE = 0.03 s, flip angle = 70°, voxel size = 2.4 × 2.4 × 2.5 mm, inter-slice gap = 0.5 mm, field of view (FOV) = 192 × 172.8 × 107.5. Participants completed 1 to 4 scans of 240 volume each (average scan time = 710.4 second per person). During the resting state scan, participants were instructed to relax but remain awake. Sighted participants wore light-excluding blindfolds to equalize the light conditions across the groups during the scans.

Removed in the revised version: Infants (dHCP)

Earlier version: Anatomical and functional images were collected on a 3T Phillips scanner at the Evelina Newborn Imaging Centre, St Thomas’ Hospital, London, UK. A dedicated neonatal imaging 219 system including a neonatal 32-channel phased-array head coil was used. T2w multi-slice fast spin-echo images were acquired with in-plane resolution 0.8x0.8 mm2 and 1.6 mm slices overlapped by 0.8 mm (TR = 12000 ms, TE = 156 ms, SENSE factor 2.11 axial and 2.6 sagittal). In infants, T2w images were used as the anatomical image because the brain anatomy is more clearly in T2w than in T1w images. Fifteen minutes of resting state fMRI data were collected using a used multiband 9x accelerated echo-planar imaging (TR = 392 ms, TE = 38 ms, 2300 volumes, with an acquired resolution of 2.15 mm isotropic). Single-band reference scans were acquired with bandwidth-matched readout, along with additional spin-echo acquisitions with both AP/PA fold-over encoding directions.

Revised version: Infants (dHCP) Anatomical and functional images were collected on a 3T Phillips scanner at the Evelina Newborn Imaging Centre, St Thomas’ Hospital, London, UK. A dedicated neonatal imaging 219 system including a neonatal 32-channel phased-array head coil was used. T2w multi-slice fast spin-echo images were acquired with in-plane resolution 0.8×0.8 mm2 and 1.6 mm slices overlapped by 0.8 mm (TR = 12000 ms, TE = 156 ms, SENSE factor 2.11 axial and 2.6 sagittal). In infants, T2w images were used as the anatomical image because the brain anatomy is more clearly in T2w than in T1w images. Fifteen minutes of resting state fMRI data were collected using a used multiband 9x accelerated echo-planar imaging (TR = 392 ms, TE = 38 ms, 2300 volumes, with an acquired resolution of 2.15 mm isotropic). Single-band reference scans were acquired with bandwidth-matched readout, along with additional spin-echo acquisitions with both AP/PA fold-over encoding directions.

Data analysis

Resting state data were preprocessed using FSL version 5.0.9 (Smith et al., 2004), DPABI version 6.1 (Yan et al., 2016), FreeSurfer (Dale et al., 1999), and in-house code (https://github.com/NPDL/Resting-state_dHCP). The functional data for all groups were linearly detrended and low-pass filtered (0.08 Hz).

For adults, functional images were registered to the T1-weighted structural images, motion corrected using MCFLIRT (Jenkinson et al., 2002), and temporally high-pass filtering (150 s). No subject had excessive head movement (> 2mm) or rotation (> 2°) at any timepoint. Resting state data are known to include artifacts related to physiological fluctuations such as cardiac pulsations and respiratory-induced modulation of the main magnetic field. A component-based method, CompCor (Behzadi et al., 2007), was therefore used to control for these artifacts. Particularly, following the procedure described in Whitfield-Gabrieli et al., nuisance signals were extracted from 2-voxel eroded masks of spinal fluid (CSF) and white matter (WM), and the first 5 principal components analysis (PCA) components derived from these signals was regressed out from the processed BOLD time series (Whitfield-Gabrieli & Nieto-Castanon, 2012). In addition, a scrubbing procedure was applied to further reduce the effect of motion on functional connectivity measures (Power et al., 2012, 2014). Frames with root mean square intensity difference exceeding 1.5 interquartile range above the 75th centile, after motion and distortion correction, were censored as outliers.

Earlier version: The infants resting state functional data were pre-processed by the dHCP group using the project’s in-house pipeline (Fitzgibbon et al., 2020), This pipeline uses a spatial independent component analysis (ICA) denoising step to minimize artifact due to multi-band artefact, residual head-movement, arteries, sagittal sinus, CSF pulsation. For infants, ICA denoising is preferable to using CSF/WM regressors. Because it is challenging to accurately define anatomical boundaries of CSF/WM due to the low imaging resolution comparing with the brain size and the severe partial-volume effect in the neonate (Fitzgibbon et al., 2020). Like in the adults, frames with root mean square intensity difference exceeding 1.5 interquartile range above the 75th centile, after motion and distortion correction, were considered as motion outliers. Out from the 2300 frames, a subset of continuous 1600 with minimum number of motion outliers was kept for each subject. Motion outliers were censored from the subset of continuous 1600, and a subject was excluded from further analyses when the number of outlier exceeded 160 (10% of the continues subset) (Hu et al., 2022).

Revised version: The infants resting state functional data were pre-processed by the dHCP group using the project’s in-house pipeline (Fitzgibbon et al., 2020), This pipeline uses a spatial independent component analysis (ICA) denoising step to minimize artifact due to multi-band artefact, residual head-movement, arteries, sagittal sinus, CSF pulsation. For infants, ICA denoising is preferable to using CSF/WM regressors. Because it is challenging to accurately define anatomical boundaries of CSF/WM due to the low imaging resolution comparing with the brain size and the severe partial-volume effect in the neonate (Fitzgibbon et al., 2020). Like in the adults, frames with root mean square intensity difference exceeding 1.5 interquartile range above the 75th centile, after motion and distortion correction, were considered as motion outliers. Out from the 2300 frames, a subset of continuous 1600 with minimum number of motion outliers was kept for each subject. Motion outliers were censored from the subset of continuous 1600, and a subject was excluded from further analyses when the number of outlier exceeded 160 (10% of the continues subset) (Hu et al., 2022). While infant connectivity estimates may be less robust at the individual level compared to adults due to shorter scan durations and higher motion, our cohort’s large sample size (n=475) and rigorous motion censoring mitigate these limitations for group-level analyses. Substantial differences between the groups exist in this study, including the number of subjects, brain sizes, imaging parameters, and data preprocessing, all of which are likely to have an impact on the overall signal quality. To address this concern, we compared the split-half noise ceiling across the groups (infants, sighted adults, and blind adults). For each participant, we first split the rs-fMRI time series into two halves, then calculated the ROI-wise rsFC pattern from the two splits. The split-half noise ceiling was estimated according to Lage-Castellanos et al (Lage-Castellanos et al., 2019). The noise ceilings of the three groups (infants: 0.90 ± 0.056, blind adults: 0.88 ± 0.041, sighted adults: 0.90 ± 0.055) showed no significant difference (One-way ANOVA, F (2,552) = 2.348, p = 0.097). Therefore, overall signal quality is unlikely to impact our results.

For both groups of adult and infants, we performed a temporal low-pass filter (0.08 Hz low-pass cutoff) and a linear detrending. ROI-to-ROI connectivity was calculated using Pearson’s correlation between ROI-averaged BOLD timeseries (ROI definition see below). The All t-tests and F-tests are two-sided. The comparison of correlation coefficients was done using cocor software package and Pearson and Filon’s z (Diedenhofen & Musch, 2015; Pearson & Filon, 1898).

ROI definition

Earlier version: ROIs in the frontal and occipital cortices were defined from separate task-based fMRI experiments with blind and sighted adults (Kanjlia et al., 2016, 2021; Lane et al., 2015). Three separate experiments were conducted with the same group of blind and sighted subjects (sighted n=18; blind n=23). The language ROIs in the occipital and frontal cortices were identified by sentence > nonwords contrast in an auditory language comprehension task (Lane et al., 2015). The math ROIs were identified by math > sentence contrast in an auditory task where participants judged equivalence of pairs of math equations and pairs of sentences (Kanjlia et al., 2016). The executive function ROIs were identified by no-go > frequent go contrast in an auditory go/no-go task with non-verbal sounds (Kanjlia et al., 2021). The occipital ROIs were defined based on group comparisons blind > sighted in a whole-cortex analysis. For example, the occipital language ROI were defined as the cluster that responded more to auditory sentence than auditory nonwords conditions in blind, relative to sighted, in a whole-cortex analysis. All three occipital ROIs were defined in the right hemisphere. Left-hemisphere occipital ROIs were created by flipping the right-hemisphere ROIs to the left hemisphere. The frontal ROIs were defined based on a whole-cortex analysis which combined all blind and sighted adult data. For example, the frontal language ROI was defined as responded more auditory sentence than auditory nonwords conditions across all blind and sighted subjects, constrained to the prefrontal cortex. For frontal ROIs, the language ROI was defined in the left, and the math and executive function ROI were defined in the right hemisphere, then flip to the other hemisphere.

Revised version: Frontal and secondary visual ROIs were defined functionally based on data from a separate task-based fMRI experiments with blind and sighted adults (Kanjlia et al., 2016, 2021; Lane et al., 2015). Three separate experiments were conducted with the same group of blind and sighted subjects (sighted n=18; blind n=23). The language ROIs in the occipital and frontal cortices were identified by sentence > nonwords contrast in an auditory language comprehension task (Lane et al., 2015). The math ROIs were identified by math > sentence contrast in an auditory task where participants judged equivalence of pairs of math equations and pairs of sentences (Kanjlia et al., 2016). The executive function ROIs were identified by no-go > frequent go contrast in an auditory go/no-go task with non-verbal sounds (Kanjlia et al., 2021). All ROI files are available at openICPSR (https://doi.org/10.3886/E198832V1).

Added in the revised version: Occipital secondary ‘visual’ ROIs were defined based on group comparisons blind > sighted in a whole-cortex analysis (Figure 1, Figure S5.) The occipital language ROI was defined as the cluster that responded more to auditory sentence than auditory nonwords conditions in blind, relative to sighted, in a whole-cortex analysis, likewise the occipital math ROI was defined as math>sentences, blind>sighted interaction and the occipital executive ROI as no-go > frequent go, blind>sighted. All three occipital ROIs were defined in the right hemisphere. Left-hemisphere occipital ROIs were created by flipping the right-hemisphere ROIs to the left hemisphere. Each functional ROI spans multiple anatomical regions and together the secondary visual ROIs tile large portions of lateral occipital, occipito-temporal, dorsal occipital and occipito-parietal cortices. In sighted people, the secondary visual occipital ROIs include the anatomical locations of functional regions such as motion area V5/MT+, the lateral occipital complex (LO), category specific ventral occipitotemporal cortices and dorsally, V3a and V4v. The occipital ROI also covers the middle of the ventral temporal lobe. Dorsally, it extended to the intraparietal sulcus and superior parietal lobule.

Added in the revised version: The frontal PFC ROIs were defined functionally, based on a whole-cortex analysis which combined all blind and sighted adult data. The frontal language ROI was defined as responded more auditory sentence than auditory nonwords conditions across all blind and sighted subjects, constrained to the prefrontal cortex. Likewise, math responsive PFC was defined as math>sentences and executive no-go > frequent go. For frontal ROIs, the language ROI was defined in the left, and the math and executive function ROI were defined in the right hemisphere, then flip to the other hemisphere.

Earlier version: The V1 was defined from a previously published anatomical surface-based atlas (PALS-B12) (Van Essen, 2005). The primary somatosensory and motor cortex (S1/M1) ROI was selected as the area that responds more to the go than no-go trials in the auditory go/no-go task across both blind and sighted groups, constrained to the hand area in S1/M1 search-space from neurosynth.org (term “hand movements”) (Kanjlia et al., 2021). The primary auditory cortex (A1) ROI was defined as the transverse temporal portion of a gyral-based atlas (Desikan et al., 2006; Morosan et al., 2001).

Revised version: The V1 ROI was defined from a previously published anatomical surface-based atlas (PALS-B12) (Van Essen, 2005). The primary somatosensory and motor cortex (S1/M1) ROI was selected as the area that responds more to the go than no-go trials in the auditory go/no-go task across both blind and sighted groups, constrained to the hand area in S1/M1 search-space from neurosynth.org (term “hand movements”) (Kanjlia et al., 2021). The primary auditory cortex (A1) ROI was defined as the transverse temporal portion of a gyral-based atlas (Desikan et al., 2006; Morosan et al., 2001).

Earlier version: All the ROIs were defined in standard space. For infants, the ROIs were subsequently transformed into each subject’s native space using a two-step approach. First, the ROIs were converted from the adult’s MNI space into the 40-week dHCP template (Bozek et al., 2018). ANTS, previously shown to be effective in pediatric studies (Avants et al., 2014; Cabral et al., 2022; Jain et al., 2012; Lawson et al., 2013), was utilized to estimate the deformation field between these two spaces. In this step, the infant’s scalp and cerebellum were masked, as these structures in the infant brain greatly differ from those in the adult and can introduce bias into the registration process, as outlined in a study by Cabral et al (Cabral et al., 2022). Secondly, the ROIs were further transformed from the 40-week space into individual’s native spaces, employing the deformation field provided by the dHCP group. Nearest neighbor interpolation was applied in both steps (see examples of ROI in individual infant brains in the Supplementary Figure S15). For adults, the ROIs were also transformed into each subject’s native space, employing the deformation field estimated by FreeSurfer. For both adults and infants, any overlapping voxels between ROIs were removed and not counted toward any ROIs.

Revised version: All the ROIs were defined in standard space and then transformed into each subject’s native space. For adults this was done by employing the deformation field estimated by FreeSurfer. For infants, the ROIs were transformed into each subject’s native space using a two-step approach. First, the ROIs were converted from the adult’s MNI space into the 40-week dHCP template (Bozek et al., 2018). ANTS, previously shown to be effective in pediatric studies (Avants et al., 2014; Cabral et al., 2022; Jain et al., 2012; Lawson et al., 2013), was utilized to estimate the deformation field between these two spaces. In this step, the infant’s scalp and cerebellum were masked, as these structures in the infant brain greatly differ from those in the adult and can introduce bias into the registration process, as outlined in a study by Cabral et al (Cabral et al., 2022). Secondly, the ROIs were further transformed from the 40-week template space into each individual’s native spaces, employing the deformation field provided by the dHCP group. Nearest neighbor interpolation was applied in both steps (Examples of the resulting ROI alignment on individual brains are shown in Supplementary Figure S12). For both adults and infants, any overlapping voxels between ROIs were removed and not counted toward any ROIs.

Removed in the revised version: Data availability

Removed in the revised version: Neonate data were from the second and third release of the Developing Human Connectome Project (https://www.developingconnectome.org). The de-identified blind and sighted adults’ data will be posted on Vivli.org upon publication of the current manuscript.

Removed in the revised version: Supporting information